Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment Week 2

Welcome to the second weekly assignment! Every week we'll give you an assignment that covers the materiel from the previous week. Again, these assignment are not obligatory but are highly recommended.

I'm sure all of you know that the Puffin presents a remarkable combination of elegance and cuteness that is unparalleled in the animal kingdom. What you may not know is that [the Puffin is endangered](https://www.nytimes.com/interactive/2018/08/29/climate/puffins-dwindling-iceland.html). In fact, [global populations have plummeted in recent years](https://www.smithsonianmag.com/science-nature/uncertain-future-puffin-dinner-180961829/) due to several factors such as climate change and hunting (not helped by that the Puffins are, allegedly, delicious). The first step to finding a solution is to understand the problem. Hence, for this assignment we're going to build a system for simulating the Puffin population. We're going to do this for an imaginary population and we'll make some simplifying assumptions but the general idea is applicable to many real-world scenarios.

## Problem 1
One of the ways of protecting the population is to decide the hunting quota, i.e., how many Puffins hunters are allowed to hunt. Write a function `quota1(nadults)` that takes the number of adult Puffins as its only argument and returns the hunting quota. Assume that the total quota is 30% of all adults.

### Python Hint: Function definitions
```python
def twice(a):     # define function named 'twice', taking 'a' as argument
    return a + a  # compute a+a and return the value to the caller
```

<!-- grade_id: cell-e1f7c8f2bae79f2e, points: 0, grade: false, solution: true -->
```python
def quota1(nadults):
    a=nadults*0.3
    return a
```

<!-- grade_id: cell-ede5c1845eb38473, points: 1, grade: true, solution: false -->
```python
assert quota1(1000) == 300, quota1(1000)
assert quota1(2500) == 750, quota1(2500)
print('tests pass')
```

Output:
```
tests pass
```

## Problem 2
It turns out there's a problem with our hunting policy! In some years there are fewer sand eels, which are an important part of the Puffin diet. As a result fewer Puffin chicks are born in those years. Unless we take that into account we risk wiping out the population. Write a new function `quota2(nchicks, nadults)` that implements the following policy:
* If the number of chicks is less than 20% of the number of adults, the total hunting quota is 10% of the number of adults.
* Otherwise the total hunting quota is 30% of the number of adults.

#### Python Hint: if statements
```python
if 2+2 < 5: # run the indented code if condition is true
    print("true")
    # do more stuff if you want, e.g. 'return' or 'a = b', or something
```
#### Python Hint: comparison
Less than: `a < b`, less-or-equal: `a <= b`, equal: `a == b`, not equal: `a != b`, greater-or-equal: `a >= b`, greater than: `a > b`. `a` and `b` can be any expressions, not just variables: `a*3 < 1/b`.

<!-- grade_id: cell-45eb2240cd28593e, points: 0, grade: false, solution: true -->
```python
def quota2(nchicks, nadults):
    proc=nchicks/nadults
    if proc < 0.2:
        return nadults*0.1
    return nadults*0.3
```

<!-- grade_id: cell-5af8504b6f4e0541, points: 1, grade: true, solution: false -->
```python
assert quota2(10, 1000) == 100, quota2(10, 1000)
assert quota2(20, 100) == 30, quota2(20, 100)
assert quota2(19, 100) == 10, quota2(19, 100)
assert quota2(4000, 10000) == 3000, quota2(4000, 10000)
print('tests pass')
```

Output:
```
tests pass
```

## Problem 3
It turns out that even the above policy is problematic in years with especially few sand eels. You decide that if the number of chicks is less than 10% of the number of adults no hunting whatsoever is allowed. Write a new function `quota2(nchicks, nadults)` that implements the following policy
* If the number of chicks is less than 10% of the number of adults, no hunting is allowed.
* If the number of chicks is at least 10% but less than 20% of the number of adults, the total hunting quota is 10% of the number of adults.
* Otherwise the total hunting quota is 30% of the number of adults.

#### Python Hint: returning mulitple alternatives
You can combine `if` and `return` to return different values in different cases:
```python
def howMany(a):
    if a == 0:
        return "none"
    if a == 1:
        return "one"
    if a < 3:
        return "some"
    return "many" # this one will happen if none of the above match
```


<!-- grade_id: cell-b72ca823db5436a5, points: 0, grade: false, solution: true -->
```python
def quota3(nchicks, nadults):
    proc=nchicks/nadults
    if proc < 0.1:
        return 0
    if proc < 0.2:
        return 0.1*nadults
    return nadults*0.3
        
```

<!-- grade_id: cell-6717fa85be0bf312, points: 1, grade: true, solution: false -->
```python
assert quota3(10, 1000) == 0, quota3(10, 1000)
assert quota3(20, 100) == 30, quota3(20, 100)
assert quota3(19, 100) == 10, quota3(19, 100)
assert quota3(4000, 10000) == 3000, quota3(4000, 10000)
print('tests pass')
```

Output:
```
tests pass
```

## Problem 4
The latest policy is looking promising. There are still many aspects in the life of the Puffins we've not included but what we have is a good start. For this problem we're going to write a function that simulates one year of the Puffin population. Write a function `simulate1(nage0, nage1, nage2, nadults)` that takes as arguments the number of Puffins of age 0 (newborns), 1 and 2 years as well as the number of adult Puffins at the beginning of the year and returns the number of adult Puffins at the end of the year. 
* Puffins become adults at age 3, i.e., the number of adults increases by `nage2` by the end of the year.
* You should compute the hunting quota using the `quota3` function.
* The hunting quota should be computed from the number of Puffins at the beginning of the year. 
* The `nchicks` argument for `quote3` is sum of `nage0`, `nage1` and `nage2`.

<!-- grade_id: cell-2e0d2298149a955a, points: 0, grade: false, solution: true -->
```python
def simulate1(nage0, nage1, nage2, nadults):
    '''
    returns the number of adult Puffins at the end of the year 
    based on the number of Puffins at the beginning of the year.
    
    args:
    
    nage0: number of Puffins at age 0 (newborns)
    
    nage1, nage2: number of Puffins of age 1, 2.
    
    nadult: number of adult Puffins.
    
    returns: the number of adult Puffins at the end of the year.
    '''
    add=nage0+nage1+nage2
    hunt=quota3(add,nadults)
    return nadults-hunt+nage2   
```

<!-- grade_id: cell-5727e7a33dbc083f, points: 2, grade: true, solution: false -->
```python
assert simulate1(10, 11, 12, 30) == 33, simulate1(10, 11, 12, 30)
assert simulate1(1, 2, 3, 30) == 24, simulate1(1, 2, 3, 30)
assert simulate1(100, 2, 3, 400) == 283, simulate1(100, 2, 3, 400)
assert simulate1(100, 2, 3, 4000) == 4003, simulate1(100, 2, 3, 4000)
print('tests pass')
```

Output:
```
tests pass
```

## Problem 5
Puffins (unfortunalely) also die due to natural causes. Write a new function `deathrate1(npuffins, age)` that takes as arguments the number of Puffins and their age and returns the estimated number of Puffins that die in each year due to natural causes. Use `age=3` for all adults. Use the following values for the death rate:
* 20% of newborn Puffins, i.e., `age=0`, die in each year.
* 5% of Puffins of age 1, i.e., `age=1`, die in each year.
* 1% of Puffins of age 2, i.e., `age=2`, die in each year.
* 0.2% of adult Puffins, i.e., `age=3`, die in each year.

<!-- grade_id: cell-268fa507e371dc9c, points: 0, grade: false, solution: true -->
```python
def deathrate1(npuffins, age):
    '''return the estimated number of Puffins of a given age 
    that die in each year due to natural causes.    
    '''
    if age==0:
        return npuffins*0.2
    if age==1:
        return npuffins*0.05
    if age==2:
        return npuffins*0.01
    if age==3:
        return npuffins*0.002
    raise ValueError('age must be 0 an 3')
```

<!-- grade_id: cell-451119b839f1dcaf, points: 2, grade: true, solution: false -->
```python
assert deathrate1(25, 0) == 5
assert deathrate1(100, 1) == 5
assert deathrate1(1000, 2) == 10
assert deathrate1(23443, 3) == 46.886
assert deathrate1(100, 1) == 5
print('tests pass')
```

Output:
```
tests pass
```

## Problem 5
Well, it turns out that if there are too many Puffins there's not enough food to go around and the death rate increase for Puffins of all ages. Denote by `ntotal` the total number of Puffins of all ages, i.e., `ntotal=nage0+nage1+nage2+nadults`. The Puffin scientists (Puffintists?) have figured out that the death rates in problem 4 are correct as long as the total number of Puffins is at most 10000, i.e., `ntotal<=10000`. If the number of Puffins is above this number the death rate increase by a factor `ntotal/10000` (this means you need to multiply the previous deathrate by `ntotal/10000` if `ntotal>10000`. Write a new function `deathrate2(npuffins, age, ntotal)` that takes this into account.

<!-- grade_id: cell-cc81067e40f4f635, points: 0, grade: false, solution: true -->
```python
def deathrate2(npuffins, age, ntotal):
    '''return the estimated number of Puffins of a given age 
    that die in each year due to natural causes. takes into account 
    food scarcity causes by the population growing too large.
    
    '''
    if ntotal<=10000:
        return deathrate1(npuffins,age)
    return deathrate1(npuffins,age)*ntotal/10000
```

<!-- grade_id: cell-f61b3acfacd73561, points: 2, grade: true, solution: false -->
```python
tol = 0.01
assert deathrate2(100, 0, 1000) == 20
assert deathrate2(100, 0, 20000) == 40
assert deathrate2(25, 0, 25) == 5
assert deathrate2(100, 1, 100) == 5
assert deathrate2(1000, 2, 1000) == 10
assert abs(deathrate2(23443, 3, 23445) - 109.924227) < tol
assert deathrate2(100, 1, 100) == 5
print('tests pass')
```

Output:
```
tests pass
```

## Problem 6
There's a bug in the deathrate2 function. Specifically, it's possible for the function to report a total number of deaths that is larger than the number of Puffins given. Write a new function `deathrate3(npuffins, age, ntotal)` that never returns a value larger that `npuffins`. You can do this with only one additional line of code.

<!-- grade_id: cell-cc1d1e5eb925f8e2, points: 0, grade: false, solution: true -->
```python
def deathrate3(npuffins, age, ntotal):
    '''return the estimated number of Puffins of a given age 
    that die in each year due to natural causes. takes into account 
    food scarcity causes by the population growing too large.
    
    '''
    if deathrate2(npuffins,age,ntotal)>npuffins:
        return npuffins
    return deathrate2(npuffins,age,ntotal)
```

<!-- grade_id: cell-39bcf2706e6ab203, points: 1, grade: true, solution: false -->
```python
tol = 0.01
assert deathrate3(100, 0, 1000) == 20
assert deathrate3(100, 0, 20000) == 40
assert deathrate3(25, 0, 25) == 5
assert deathrate3(100, 1, 100) == 5
assert deathrate3(1000, 2, 1000) == 10
assert abs(deathrate3(23443, 3, 23445) - 109.924227) < tol
assert deathrate3(100, 1, 100) == 5
assert deathrate3(10, 0, 60000) == 10
assert deathrate3(11, 1, 10000*21) == 11
assert deathrate3(12, 2, 10000*101) == 12
assert deathrate3(14, 3, 10000*501) == 14
print('tests pass')
```

Output:
```
tests pass
```

## Problem 7
For the final problem we're going to include the deathrate in our simulation of the Puffin population. Write a new function `simulate2(nage0, nage1, nage2, nadults)`. This function should work like the one you wrote for problem 4 except that it should also include the deathrate.
* Use the `deathrate3` function to compute the death rate. The death rate is computed based on the number of Puffins after the hunting quota is subtracted.
* The hunting quota (computed using `quota3`) should be computed based on the number of Puffins at the beginning of the year, i.e., before the death rate is applied and before `nage2` is added to the number of adults. Adding `nage2` to `nadults` should be the final step.

<!-- grade_id: cell-f1b454b2c2f9acb6, points: 0, grade: false, solution: true -->
```python
def simulate2(nage0, nage1, nage2, nadults):
    '''
    returns the number of adult Puffins at the end of the year 
    based on the number of Puffins at the beginning of the year.
    takes death rate into account.
    
    args:
    
    nage0: number of Puffins at age 0 (newborns)
    
    nage1, nage2: number of Puffins of age 1, 2.
    
    nadult: number of adult Puffins.
    
    returns: the number of adult Puffins at the end of the year.
    '''
    # smiertelnosc dla 2 gen.
    
    nchicks=nage0+nage1+nage2
    nadults-=quota3(nchicks,nadults)
    
    total=nchicks+nadults
    nage2 -= deathrate3(nage2, 2, total)
    print('nage2', nage2)
    nadults -= deathrate3(nadults, 3, total)
    nadults += nage2
    return nadults
    
    '''
    npuffins=nage0+nage1+nage2
    
    gen2=deathrate3(nage2,2,nadults-quota3(npuffins,nadults))
    gen3=deathrate3(nadults-quota3(npuffins,nadults),3,nadults-quota3(npuffins,nadults))
    return nadults-quota3(npuffins,nadults)-gen3+nage2-gen2
    '''
    
```

<!-- grade_id: cell-f7fedc0dd4800f6a, points: 3, grade: true, solution: false -->
```python
tol = 0.01
assert abs(simulate2(10, 11, 12, 30) - 32.838) < tol, simulate2(10, 11, 12, 30)
assert abs(simulate2(20, 10, 40, 302) - 250.5772) < tol, simulate2(20, 10, 40, 302)
assert abs(simulate2(100, 2, 3, 400) - 282.41) < tol, simulate2(100, 2, 3, 400)
assert abs(simulate2(100, 2, 3, 4000) - 3994.97) < tol, simulate2(100, 2, 3, 4000)
assert abs(simulate2(100, 20, 333, 40000) - 39995.905151) < tol, simulate2(100, 20, 333, 40000)
assert abs(simulate2(133333, 2333, 3000, 100000) - 69452.678) < tol, simulate2(133333, 2333, 3000, 100000)
print('tests pass')
```

Output:
```
nage2 11.88
nage2 39.6
nage2 2.97
nage2 2.97
nage2 319.529151
nage2 2374.002
tests pass
```

## Getting Ahead
Congratulations on completing the assignment! Nice work :)

This assignment will ultimately do little to alleviate the problems the Puffins of today face. However, while we've made several simplifying assumptions the general method used above is applicable to a wide range of scenarios in the real world. Scientists and decision makers rely on models not unlike this one in every field imaginable. The beauty of programming (and Python is an especially good example) is that it gives you access to the tools you need to reason about the world. Even a simple program like this can offer a surprising amount of insight into complex problems.

One of the shortcomings of our system is that it only simulates a single year. To design policy we need to simulate the population over many years. Hence, for the getting ahead part we're goubg to write a new function `nsimulate(years, nage0, nage1, nage2, nadults)` that simulautes `years` years of the population and returns the number of adult Puffins at the end of this period. The function should take as arguments the number of years to simulate and the number of Puffins at the start of the first year. Each year should behave like in `simulate2` that we wrote for problem 7. To do this you'll need to [learn about for loops](https://www.youtube.com/watch?v=9LgyKiq_hU0). Remember that you need to update the number of Puffins of each age for each cycle of the simulation. Assume that each adult Puffin (as counted by the end of the year, i.e., after the death rate and hunting quota has been subtracted and after `nage2` has been added to `nadults`) has 0.4 chicks in each year, i.e., `nage0` for the next year is computed as `nage0 = nadults * 0.4`, where `nadults` is the number of adults at the end of the current year.

This part is meant as preparation for the coming course material and includes materiel not yet covered in the course. Then again, you never have to catch up if you're always getting ahead.

<!-- grade_id: cell-1b7041bc1bc8e232, points: 0, grade: false, solution: true -->
```python
def nsimulate(years, nage0, nage1, nage2, nadults):
    '''
    returns the number of adult Puffins at the end of years years 
    based on the number of Puffins at the beginning of the year 0.
    takes death rate into account.
    
    args:
    
    years: the number of years to simulate.
    
    nage0: number of Puffins at age 0 (newborns) at the beginning of year 0.
    
    nage1, nage2: number of Puffins of age 1, 2 at the beginning of year 0.
    
    nadult: number of adult Puffins at the beginning of year 0.
    
    returns: the number of adult Puffins after years years.
    '''
     # YOUR CODE HERE
    
    for _ in range(years):
         
        nchicks=nage0+nage1+nage2
        nadults-=quota3(nchicks,nadults)
        total=nchicks+nadults
        nage2 -= deathrate3(nage2, 2, total)
        nage0 -= deathrate3(nage0, 0, total)
        nage1 -= deathrate3(nage1, 1, total)
        nadults -= deathrate3(nadults, 3, total)
        
        nadults += nage2
        nage2 = nage1
        nage1 = nage0
        nage0 = nadults*0.4
        
    return nadults
   
 
  

```

<!-- grade_id: cell-e0f4372b3388371b, points: 2, grade: true, solution: false -->
```python
tol = 0.01
assert abs(nsimulate(10, 10, 11, 15, 32) - 34.26521955455017) < tol, nsimulate(10, 10, 11, 15, 32)
assert abs(nsimulate(11, 1, 2, 3, 30) - 16.54276141994387) < tol, nsimulate(11, 1, 2, 3, 30)
assert abs(nsimulate(12, 100, 2, 3, 400) - 223.68173879099925) < tol, nsimulate(12, 100, 2, 3, 400)
assert abs(nsimulate(5, 100, 2, 3, 4000) - 2669.359577430396) < tol, nsimulate(5, 100, 2, 3, 4000)
assert abs(nsimulate(4, 100, 20, 333, 40000) - 15011.789929962835) < tol, nsimulate(4, 100, 20, 333, 40000)
assert abs(nsimulate(33, 133333, 2333, 3000, 100000) - 6444.8841810316235) < tol, nsimulate(33, 133333, 2333, 3000, 100000)
print('tests pass')
```

Output:
```
tests pass
```

If you did it, nice work!

Feel free to keep playing around with creating new hunting policies. Can you create a better or more realistic policy than the one we've considered? Try experimenting with this system by writing a program that you can run locally on your own computer. Otherwise no more content for this week!