Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment Week 3

Welcome to the third weekly assignment! Every week we'll give you an assignment that covers the materiel from the previous week. Again, these assignment are not obligatory but are highly recommended.

You've landed a job at the [Lon Lon Ranch](http://zelda.wikia.com/wiki/Lon_Lon_Ranch). The ranch stock a wide range of items and it's your job to keep track of what it's all worth. Specifically, you've been tasked with writing a function that given an item name, its age and (optionally) further information returns the value of the item. Items (such as milk) may change in value over time.

This assignment differs from the previous one in that you'll be working on the same function throughout the whole assignment. Initially the function will be short. However, as more items and rules are added you'll need to add more things to the function. Keep in mind that adding additional functions that are called from the main function is a good strategy to keep each function simple and easy to reason about. Furthermore, you'll need to take some care to figure out what's wrong when the tests don't work.

## General Guidelines
These are general guidelines that will (hopefully) help you preserve your sanity through your programming endeavours.
* Set aside time to clean up your code. The number that gets bandied around on the internet is that the best programmers spend roughly half their time cleaning up and rewriting their code, i.e., you are halfway done when you get it working.
* The sooner you start programming the longer it takes, i.e., make sure you understand the problem and have an idea for how to solve it before you get started.
* Make sure you understand the current assignment before moving on. Building on bedrock is better than building on sand.
* You should be spending a significant amount of time on [Stack Overflow](https://stackoverflow.com/) and [Google](https://duckduckgo.com/?q=how+to+python). Even experienced programmers do this.
* [PEP8](https://www.python.org/dev/peps/pep-0008/) is the official Python style guilde. Following these guidelines makes sure your code is neat and easy to read.

## Function
You'll be working in the code cell below throughout the whole assignment. As we're adding more items and rules you should make modifications as needed to this function. **Make sure to re-run the whole notebook by pressing the double play button at the top (restart the kernel, then re-run the whole notebook) whenever you make any changes.** This is due to two reasons:
* You need to rerun a cell with code in it whenever you make changes for the changes to take effect.
* You should re-run all tests whenever you change something to make sure your change didn't break an earlier test.
Re-running the whole notebook takes care of this automatically.

The [signature](https://en.wikipedia.org/wiki/Type_signature) of the function you'll be working on is `get_value(item_name, item_age, item_details)`. The parameters are:
* item_name: A string containing the item name, e.g., "milk".
* item_age: The age of the item measured in days.
* item_details: A string containing further details about the item, e.g., "lon lon". For items with no details this argument should be the empty string (`""`).

When learning programming as you do now it's common to look back on the code written by past you and wonder what in the world he/she was thinking. This is good because it means present you is a stronger programmer than past you. However, it also means present you has to keep correcting the things past you wrote. Therefore, it's important that you go back and clean up your code after completing each problem.

#### Python Tip: Functions, Parameters, Arguments and Return Values
Remember that when a function is called (e.g., `get_value("milk", 2*7, "")`) the *actual argument* values are computed and placed in the *formal parameter* variables (e.g., `item_name = 'milk'`, `item_age = 14`, `item_details = ''`). When the function reaches the `return` statement, the *return value* is computed and becomes the value of the function call.

<!-- grade_id: cell-dd49ba701655981b, points: 0, grade: false, solution: true -->
```python
def get_value_cheese(age):
   
        v=0
        if age<30:
            return 100
        if age>=30 and age<360:
            for i in range(30,age):
                v=v+3
                #print(age,v,i)
            return v+100
        if age>=360 and age<720:
            for i in range(30,360):
                v=v+3
            for i in range(360,age):
                v=v+4
            return v+100
        if age>=720:
            for i in range(30,360):
                v=v+3
            for i in range(360,720):
                v=v+4
            for i in range(720,age):
                v=v-10
            v=v+100
            if v<0:
                v=0
            return v
        
def get_value_whisky(age, cask_type, quality):
    
    """Any whisky younger than 360 days has value 0.
Whisky between 360 (inclusive) and 720 days has value 300 rupees.
Starting at day 720 the value of whisky increases by
1% per 30 days for "ex-sherry" cask.
2% per 30 days for "ex-bourbon" cask.
3% per 30 days for "ex-amarone" cask.
The value of the whisky remains constant between 30 day increments, i.e., you need to round down to the closest multiple of 30 days.
The overall value of the whisky should be multiplied by the quality, which is greater than 0 and at most 1."""
    v=0
    if age<360:
        return v
    if age>=360 and age <720:
        v=300
        print(quality)
        return v*quality
    if age>=720:
        if cask_type=="ex-sherry":
            v=300*pow(1.01, int(age/30-720/30))
            return v*quality
        
        if cask_type=="ex-bourbon":
            v=300*pow(1.02, int(age/30-720/30))
            return v*quality
        
        if cask_type=="ex-amarone":
            v=300*pow(1.03, int(age/30-720/30))
            return v*quality
        
    
    
    
                
def get_value(item_name, item_age, item_details):
    '''return the value of an item based on its name,  age and any further details given.
    
    args:

    item_name: the name of the item. given as a string.
  
    item_age: the age of the item measured in days.
    
    item_details: a string containing further details about the item. is the empty string ("") 
    for items with no details. 
    
    '''
   
    if item_name.lower().strip().replace(" " ,"")=="whisky":
        cask_type, quality = item_details.split(",")
        #print(quality)
        return get_value_whisky(item_age,cask_type, float(quality))
    
    if item_name.lower().strip().replace(" " ,"")=="agedcheese":
        return get_value_cheese(item_age)
    if item_name.lower().strip().replace(" " ,"")=="milk":
            if item_age<=14:
                if item_age<=7:
                    if item_details.lower().strip().replace(" ","")=="premium":
                        return 200
                    if item_details.lower().strip().replace(" ","")=="lonlon":
                        return 150
                    return 100
                if item_details.lower().strip().replace(" ","")=="premium":
                    return 150
                if item_details.lower().strip().replace(" ","")=="lonlon":
                    return 150
                return 100
            return 0
    if item_name.lower().strip().replace(" " ,"")=="emptybottle":
        return 1000
    return None
    
  

    
    
    '''
    if item_name=="milk":
        if item_age<=7 and item_details =="premium":
            return 200
        if (item_age<=14 and item_details!="lon lon") and (item_age>7 and item_details!="lon lon"):
            return 100
        if item_age<=14 and item_details=="lon lon" or item_details=="premium":
            return 150
        return 0
    return 1000
    '''
```

## Problem 1
One of the items stocked by the ranch is empty bottles. As you probably know these are some of the most valuable items in the world. Hence, the Talon, the owner of the ranch, has decided to sell them for 1000 rupees each. Update the `get_value` function so that it returns 1000 when called with these arguments:
* item_name: "empty bottle"
* item_age: Arbitrary value. Doesn't matter.
* item_details: "" (the empty string)

#### Hey, listen! if statements
```python
if 2+2 < 5: # run the indented code if condition is true
    print("true")
    # do more stuff if you want, e.g. 'return' or 'a = b', or something
```

<!-- grade_id: cell-e1f7c8f2bae79f2e, points: 2, grade: true, solution: false -->
```python
def assertEquals(a, b, tol=0):
    if isinstance(a, float) or isinstance(b, float):
        assert abs(a-b) <= tol, f'expected {b}, got {a}'
    else:
        assert a == b, f'expected {b}, got {a}'
    
def testProblem1():
    print("Problem 1:", end='')
    assertEquals(get_value("empty bottle", 0, ""), 1000)
    assertEquals(get_value("empty bottle", None, ""), 1000)
    assertEquals(get_value("empty bottle", -10, ""), 1000)
    print("    tests pass")

testProblem1()
```

Output:
```
Problem 1:    tests pass
```

## Problem 2
`"milk"` is one of the most important products produced at Lon Lon Ranch. Milk is sold for 100 rupees per bottle. However, milk spoils and becomes worthless when it gets too old! Talon has decided that milk older than 14 days cannot be sold. Update your function to return 100 when called with `"milk"` for the item name and an age at most 14 days. If the age is larger than 14 days it should return 0. The item description should be the empty string (`""`).

#### Hey, listen! multiple conditions
```python
a = 10
if a > 5 and a < 15: # run the indented code if a is greater than 5 and smaller than 15
    return "5 > a < 15"
    # do more stuff if you want, e.g. 'return' or 'a = b', or something
```

<!-- grade_id: cell-45eb2240cd28593e, points: 2, grade: true, solution: false -->
```python
def testProblem2():
    print("Problem 2:", end='')
    assertEquals(get_value("milk", 0, ""), 100)
    assertEquals(get_value("milk", 13, ""), 100)
    assertEquals(get_value("milk", 14, ""), 100)
    assertEquals(get_value("milk", 15, ""), 0)
    assertEquals(get_value("milk", 14.01, ""), 0)
    print("    tests pass")

testProblem2()
testProblem1()
```

Output:
```
Problem 2:    tests pass
Problem 1:    tests pass
```

## Problem 3
Lon Lon Ranch is a well-known brand name in Hyrule. Talon wants to exploit this fact to make more money by segmenting the market. The idea is to sell the same milk under two different labels. Update the `get_value` function to handle the following two cases. Make sure the previous tests still work after making the update.
* **item_name: "milk", item_details: ""** – the value is 100 rupees for age <= 14 and 0 otherwise.
* **item_name: "milk", item_details: "lon lon"** – the value is 150 rupees for age <= 14 and 0 otherwise.

#### Hey, listen! returning mulitple alternatives
You can combine `if` and `return` to return different values in different cases:
```python
def howMany(a):
    if a == 0:
        return "none"
    if a == 1:
        return "one"
    if a < 3:
        return "some"
    return "many" # this one will happen if none of the above match
```


<!-- grade_id: cell-b72ca823db5436a5, points: 2, grade: true, solution: false -->
```python
def testProblem3():
    print("Problem 3:", end='')
    assertEquals(get_value("milk", 0, "lon lon"), 150)
    assertEquals(get_value("milk", 13, "lon lon"), 150)
    assertEquals(get_value("milk", 14, "lon lon"), 150)
    assertEquals(get_value("milk", 15, "lon lon"), 0)
    assertEquals(get_value("milk", 14.01, "lon lon"), 0)
    assertEquals(get_value("milk", 1000, "lon lon"), 0)
    print("    tests pass")
    

testProblem3()
testProblem2()
testProblem1()

```

Output:
```
Problem 3:    tests pass
Problem 2:    tests pass
Problem 1:    tests pass
```

## Problem 4
Talon's daughter Malon has figured out a way to make milk more delicious. They want to sell this new milk, referred to as premium milk, alongside the regular and Lon Lon milk to further segment the market. Premium milk can be sold at a higher price. However, consumers also expect a higher quality so it cannot be sold for as long as the others. Specifically, its value drops already after 7 days have passed. However, until its age exceeds 14 days it can still be sold as Lon Lon milk.

You should write a new function `get_value_milk(age, item_details)` that gets called from the `get_value` function when it's given the item name `"milk"`. Write this function in the same code cell as the `get_value` function. Functions is a great way to separate functionality.

* **item_name: "milk", item_details: "premium"** – value is 200 rupees for age <= 7 and 150 rupees for age between 7 (inclusive) and 14 days.

#### Hey, listen! if and elif
Use `if` with `elif` and/or `else` to select one of several alternatives:
```python
s = "foo"
v = 10
if s == "foo" and v < 8: # run the indented code if v is smaller than 8
    print("v is less than 8")
     # do more stuff if you want, e.g. 'return' or 'a = b', or something
elif s == "foo" and v < 12: # run if v is not smaller than 8 but is smaller than 12
    print("v is less than 12 and not smaller than 8")
    # do more stuff if you want, e.g. 'return' or 'a = b', or something
```

In general, it is often better to use `elif`/`else` even if all your `if`s end with `return`, since that makes it clearer to the reader that we are selecting one of several alternatives:
```python
    # ....
    if a == b:
        return "equal"
    else:  
        return "different"
```

<!-- grade_id: cell-2e0d2298149a955a, points: 2, grade: true, solution: false -->
```python
def testProblem4():
    print("Problem 4:", end='')
    assertEquals(get_value("milk", 0, "premium"), 200)
    assertEquals(get_value("milk", -1, "premium"), 200)
    assertEquals(get_value("milk", 6, "premium"), 200)
    assertEquals(get_value("milk", 7, "premium"), 200)
    assertEquals(get_value("milk", 8, "premium"), 150)
    assertEquals(get_value("milk", 7.01, "premium"), 150)
    assertEquals(get_value("milk", 1000, "premium"), 0)
    assertEquals(get_value("milk", 15, "premium"), 0)
    assertEquals(get_value("milk", 14.01, "premium"), 0)
    print("    tests pass")


testProblem4()
testProblem3()
testProblem2()
testProblem1()

```

Output:
```
Problem 4:    tests pass
Problem 3:    tests pass
Problem 2:    tests pass
Problem 1:    tests pass
```

## Problem 5
Some of the employees at the ranch have trouble entering data correctly when calling the `get_value` function. You've seen all sorts of strange errors while looking over the shoulder of some of the employees!

### a) Validating inputs
What happens if you call `get_value` with a wrong item name – one that doesn't match any of your `if` statements? Try it in the cell below.

* Depending on how you implemented `get_value` you might get an error, a strange answer, or a `None` result. What do you think would be a reasonable way to handle wrong item names?
* The tests for Problems 1–4 probably haven't caught this problem at all – so we should probably also improve our tests a bit.

We'll learn good ways of dealing with errors later, but for now, it's probably best to just `return None` if `get_value` is called with an unknown `item_name`.

Update your `get_value` with a default alternative that returns `None` if you don't recognize the item name. (This might be a good place to use `else` if you have already been using `elif`.)

After you've updated `get_value` (and re-run the notebook), try it in the cell below and check that it behaves properly. (You can even try to use `assert ... == None` or `assertEquals(..., None)` and make it a proper test case!)


```python
print(get_value("wr0ng 1tem", 0 ,""))
```

Output:
```
None
```

### b) Cleaning inputs
Most of the errors made by the ranch employees can corrected easily. For example, these are the most common mistakes:
* "Lon Lon"
* " lon lon"
* "premium "
* " Premium"
* "lonlon"
* " milk"
* "Milk"
* "MILK"
* "Empty bottle"
* "empty  bottle" (double space)

Update the `get_value` function to return the correct result in all of these cases. This will require you to learn about how to modify strings. Note that this has not been covered in the lectures yet. The functions `lower()` and `replace()` described on [this site](https://docs.python.org/3/library/stdtypes.html) are what you need. Use the notebook cell below or interactive terminal on your own machine to try these things out by typing things like `"fooBAR".lower()` into it. You need to understand how they work before you attempt to add them to your code.

I recommend always comparing against a string with no spaces at all in your if statements, i.e.,  compare against `"lonlon"` instead of `"lon lon"`. This way you can strip all of the spaces and then make the comparison. It's easier to convert `" lon  lon"` into `"lonlon"` than `"lon lon"`. In addition to this change you should only need to add 2 or 4 new lines of code (depending on how you do it).

(This is a case of input *normalization* – you transform all the different variants into a single, consistent form (the ‘normal form’) with only lowercase letters and no spaces. Note that there are other potential errors that we don't catch – for instance, what happens if someone enters a negative item age? We won't deal with this here, but we'll learn to be more careful later.)


```python
print("MiLK".lower())
print("m i l k".replace(" ", ""))
print("M i L K".lower().replace(" ", ""))
```

Output:
```
milk
milk
milk
```

<!-- grade_id: cell-268fa507e371dc9c, points: 2, grade: true, solution: false -->
```python
def testProblem5():
    print("Problem 5:", end='')
    assertEquals(get_value("milk", 0, "Lon Lon"), 150)
    assertEquals(get_value("milk", 6, " premium "), 200)
    assertEquals(get_value("milk", 7, " Premium"), 200)
    assertEquals(get_value("Milk", 6, "premium"), 200)
    assertEquals(get_value("Milk", 8, "premium"), 150)
    assertEquals(get_value("milk ", 7.01, "premium"), 150)
    assertEquals(get_value("MILK", 1000, "premium "), 0)
    assertEquals(get_value("milk", 15, " Premium"), 0)
    assertEquals(get_value("milk ", 14.01, " PREMIUM  "), 0)
    assertEquals(get_value("MILK", 6, " Premium"), 200)
    assertEquals(get_value("Empty  Bottle", 1000, ""), 1000)
    assertEquals(get_value("Empty M1lk Böttle", 1000, ""), None)
    print("    tests pass")


testProblem5()
testProblem4()
testProblem3()
testProblem2()
testProblem1()

```

Output:
```
Problem 5:    tests pass
Problem 4:    tests pass
Problem 3:    tests pass
Problem 2:    tests pass
Problem 1:    tests pass
```

## Problem 6
The ranch has recently started producing `"aged cheese"`, which has the interesting property that its price increases up until a certain point at which it starts going down again. `item_details` should be the empty string. You should write a new function `get_value_cheese(age)` that gets called from the `get_value` function when it's called with the `"aged cheese"` argument. The rules for computing the value of aged cheese are:
* Aged cheese younger than 30 days has value 100 rupees.
* Aged cheese of age between 30 (inclusive) and 360 days increases in value by 3 rupees per day.
* Aged cheese of age between 360 (inclusive) and 720 days increases in value by 4 rupees per day.
* Aged cheese of age 720 days or older decreases in value by 10 rupees per day until its value reaches 0 where it remains, i.e., the value can't become negative.

For example, the value of aged cheese that is 430 days old is 100+330\*3+70\*4=1340.

#### Hey, listen! min and max
You can use the built-in functions `min()` and `max()` to get the minimum or maximum between two numbers, e.g.,
```python
a = 10
b = 100
print(min(a, b)) # prints 10
print(max(a, b)) # prints 100

c = -10
d = 0
print(min(c, d)) # prints -10
print(max(c, d)) # prints 0
```

<!-- grade_id: cell-f61b3acfacd73561, points: 3, grade: true, solution: false -->
```python
def testProblem6():
    print("Problem 6:", end='')
    assertEquals(get_value("aged cheese", 0, ""), 100)
    assertEquals(get_value("aged cheese", 1, ""), 100)
    assertEquals(get_value("aged cheese", -1, ""), 100)
    assertEquals(get_value("aged cheese", 30, ""), 100)
    assertEquals(get_value("agedcheese", 0, ""), 100)
    assertEquals(get_value("agedCheese", 1, ""), 100)
    assertEquals(get_value("Aged  Cheese", 30, ""), 100)
    assertEquals(get_value("aged cheese", 31, ""), 103)
    assertEquals(get_value("aged  cheese", 359, ""), 1087)
    assertEquals(get_value("Aged Cheese", 360, ""), 1090)
    assertEquals(get_value(" aged  cheese ", 719, ""), 2526)
    assertEquals(get_value(" aged  cheese ", 720, ""), 2530)
    assertEquals(get_value(" aged  cheese ", 800, ""), 1730)
    assertEquals(get_value(" aged  cheese ", 971, ""), 20)
    assertEquals(get_value(" aged  cheese ", 972, ""), 10)
    assertEquals(get_value(" aged  cheese ", 973, ""), 0)
    assertEquals(get_value(" aged  cheese ", 1000, ""), 0)
    assertEquals(get_value(" aged  cheese ", 430, ""), 1370)
    print("    tests pass")

testProblem6()
testProblem5()
testProblem4()
testProblem3()
testProblem2()
testProblem1()

```

Output:
```
Problem 6:    tests pass
Problem 5:    tests pass
Problem 4:    tests pass
Problem 3:    tests pass
Problem 2:    tests pass
Problem 1:    tests pass
```

## Problem 7
The final product produced at the ranch is `"whisky"`. Whisky is special in the sense that its value depends on both the cask type and the quality of the raw materials. The cask type can be one of `"ex-sherry"`, `"ex-bourbon"` and `"ex-amarone"`. Quality is expressed as a [floating point number](https://en.wikipedia.org/wiki/Floating-point_arithmetic) ([`float`](https://docs.python.org/3/library/stdtypes.html#numeric-types-int-float-complex)) between 0 and 1 (inclusive). These are given together as a single string with cask type and quality separated by a comma, e.g.,
* `"ex-sherry,0.7"`
* `"ex-sherry,0.2"`
* `"ex-amarone,0.99"`
* `"ex-bourbon,0.112"`

The cask type always comes first and the quality is always after the comma. Since whisky is stored for much longer times than cheese its value should only change every 30 days. Specifically, the rules for whisky are
* Any whisky younger than 360 days has value 0.
* Whisky between 360 (inclusive) and 720 days has value 300 rupees.
* Starting at day 720 the value of whisky increases by
    * 1% per 30 days for `"ex-sherry"` cask.
    * 2% per 30 days for `"ex-bourbon"` cask.
    * 3% per 30 days for `"ex-amarone"` cask.
* The value of the whisky remains constant between 30 day increments, i.e., you need to round down to the closest multiple of 30 days.
* The overall value of the whisky should be multiplied by the quality, which is greater than 0 and at most 1.

For example, a whiskey of age 1000 days aged on an ex-sherry cask should have value `300*pow(1.01, int(1000/30-720/30))`, where `pow(a, b)` is a raised to the power of b.

As with milk and cheese you should impement this as a separate function `get_value_whisky(age, cask_type, quality)` that is called from the `get_value` function.

#### Hey, listen!
You can split a string that is delimited by a comma into separate variables:
```python
s = "foo,bar"
# Split string s. Assigns the part before the comma to variable a and the part after to b.
a, b = s.split(",") 
print(a) # prints "foo"
print(b) # prints "bar"
```

#### Types
Remember that values have different types. For example, when you split the `str`ing `"ex-sherry,0.2"`, you get two strings (try it below, and use the `type()` function on the two variables to see that they are strings). Python won't let you treat strings like numbers, even if the text within them happen to look like numbers. Instead, you need to convert the `str` to a `float` – fortunately, there's a built-in function that does that, and its name is easy to remember: `float()`. For example:

```python
print("The string '2.0' is be different from the number 2.0:", '2.0' != 2.0)
print("The string '2.0' converted to a float is 2.0:", float('2.0') == 2.0)
```


```python
cask_type, quality = "ex-sherry,0.2".split(",")
print(quality)
```

Output:
```
0.2
```

<!-- grade_id: cell-39bcf2706e6ab203, points: 3, grade: true, solution: false -->
```python
def testProblem7():
    print("Problem 7:", end='')
    e = 1e-6 # tolerance for float comparisons, in case of rounding errors
    assertEquals(get_value("whisky", 0, "ex-sherry,1"), 0, e)
    assertEquals(get_value("whisky", 0, "ex-bourbon,1"), 0, e)
    assertEquals(get_value("Whisky", 0, "ex-amarone, 1"), 0, e)
    assertEquals(get_value(" whisky", -1, "ex-sherry,1"), 0, e)
    assertEquals(get_value(" WHISKY", -1, "EX-bourbon,1"), 0, e)
    assertEquals(get_value("whisky", -1, " ex-amarone, 1"), 0, e)
    assertEquals(get_value("whisky", 359, "ex-sherry,0.3"), 0, e)
    assertEquals(get_value("whisky", 359, "ex-bourbon, 0.4"), 0, e)
    assertEquals(get_value("whisky", 359, "ex-amarone,0.22"), 0, e)
    assertEquals(get_value("whisky", 360, "ex-sherry,1"), 300, e)
    assertEquals(get_value("whisky", 360, "ex-bourbon,1"), 300, e)
    assertEquals(get_value("whisky", 360, "ex-amarone,1"), 300, e)
    assertEquals(get_value("whisky", 360, "ex-sherry,0.5"), 150, e)
    assertEquals(get_value("whisky", 360, "ex-bourbon,0.2"), 60, e)
    assertEquals(get_value("whisky", 360, "ex-amarone,0.9"), 270, e)
    assertEquals(get_value("whisky", 390, "ex-sherry,1"), 300, e)
    assertEquals(get_value("whisky", 390, "ex-bourbon,1"), 300, e)
    assertEquals(get_value("whisky", 390, "ex-amarone,1"), 300, e)
    assertEquals(get_value("whisky", 390, "ex-sherry,0.5"), 150, e)
    assertEquals(get_value("whisky", 390, "ex-bourbon,0.2"), 60, e)
    assertEquals(get_value("whisky", 390, "ex-amarone,0.9"), 270, e)
    assertEquals(get_value("whisky", 720, "ex-sherry,1"), 300, e)
    assertEquals(get_value("whisky", 720, "ex-bourbon,1"), 300, e)
    assertEquals(get_value("whisky", 720, "ex-amarone,1"), 300, e)
    assertEquals(get_value("whisky", 720, "ex-sherry,0.1"), 30, e)
    assertEquals(get_value("whisky", 720, "ex-bourbon,0.99"), 297, e)
    assertEquals(get_value("whisky", 720, "ex-amarone,0.3"), 90, e)
    assertEquals(get_value("whisky", 1000, "ex-sherry,1"), 328.10558180530825)
    assertEquals(get_value("whisky", 1000, "ex-bourbon,1"), 358.5277705866933)
    assertEquals(get_value("whisky", 1000, "ex-amarone,1"), 391.43195514877345)
    assertEquals(get_value("whisky", 1000, "ex-sherry,0.1"), 32.810558180530826)
    assertEquals(get_value("whisky", 1000, "ex-bourbon,0.99"), 354.94249288082636)
    assertEquals(get_value("whisky", 1000, "ex-amarone,0.3"), 117.42958654463203)
    print("    tests pass")


testProblem7()
testProblem6()
testProblem5()
testProblem4()
testProblem3()
testProblem2()
testProblem1()

```

Output:
```
Problem 7:1.0
1.0
1.0
0.5
0.2
0.9
1.0
1.0
1.0
0.5
0.2
0.9
    tests pass
Problem 6:    tests pass
Problem 5:    tests pass
Problem 4:    tests pass
Problem 3:    tests pass
Problem 2:    tests pass
Problem 1:    tests pass
```

## Getting Ahead
This week we've focused on separating functionality and debugging. I'm sure you've experienced that it's not always easy to figure out what's wrong when you're dealing with lots of different cases like we've done here. Congratulations on completing it!

There's a very natural extension we can make to the code you've written. The function we've written takes arguments for a single item at a time. However, the ranch frequently processes long lists of items at a time. Write a new function `get_values(item_list)` that takes as its single argument [a list](https://docs.python.org/3/tutorial/introduction.html#lists) of lists. Each of the lists in the list contains the three fields `"item_name"`, `"item_age"` and `item_description"` at indices 0, 1 and 2, respectively. You should call the `get_value` function you've written for each element of `item_list`. To do this you need to learn about lists. I recommend looking at [this video](https://www.youtube.com/watch?v=ohCDWZgNIU0). The style is a bit cheesy but the content is good. If you're using list comprehensions you can do this in a single line of code. Using `append()` is also fine and only requires a little bit more code.

This part is meant as preparation for the coming course material and includes materiel not yet covered in the course. Then again, you never have to catch up if you're always getting ahead. This part requires doing more learning on your own. You're in very good shape if you can complete it :)

<!-- grade_id: cell-1b7041bc1bc8e232, points: 0, grade: false, solution: true -->
```python
def get_values(item_list):
    '''
    returns a list of values for each item in item_list. item_list is a list of 
    lists with fields "item_name", "item_age", "item_details at index 0, 1, 2, 
    respectively."
    '''
    values = list()
    for item in item_list:
        item_name = item[0]
        item_age = item[1]
        item_details = item[2]
        value = get_value(item_name, item_age, item_details)
        values.append(value)
    return values
```

<!-- grade_id: cell-e0f4372b3388371b, points: 0, grade: true, solution: false -->
```python
def check(a, b, tol=1e-6):
    return abs(a-b) <= tol

# define a list of items with each item represented by a dict
item_list = [
    ["empty bottle", 360, ""],
    ["milk", 7, "premium"],
    ["milk", 8, "premium"],
    ["Aged Cheese", 360, ""],
    ["aged  cheese", 719, ""],
    ["Whisky", 360, "ex-amarone,0.9"],
    ["whisky", 720, "ex-bourbon,0.99"],
    ["whisky", 1000, "ex-bourbon,0.99"],
]

# correct values for comparison
correct_values = [
    1000,
    200,
    150,
    1090,
    2526,
    270,
    297,
    354.94249288082636,
]

# call the get_values function to get the values
computed_values = get_values(item_list)

assert len(computed_values) == len(item_list), f"incorrect number of computed values: expected {len(item_list)}, got {len(computed_values)}"

# compare each computed value against the correct value
for correct_value, computed_value, item in zip(correct_values, computed_values, item_list):
    assert check(correct_value, computed_value), f"incorrect value for {item}: expected {correct_value}, got {computed_value}"
    
print("tests pass")
```

Output:
```
0.9
tests pass
```