Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Assignment Week 4

Welcome to the fourth weekly assignment! Every week we'll give you an assignment that covers the materiel from the previous week. Again, these assignment are not obligatory but are highly recommended.

I hear the l33t sp34k is all the rage with the kids (kidz?) nowadays. As with all marketing it's best to start when they're young. Hence, the UiB administration would like to improve their communication with said kids in order to in the future attract them to UiB. However, the administration is far from fluent in l33t sp34k! As a result they've hired you to write software that translates between normal speak and l33t sp34k. 

Specifically, we're going to write a pair of functions `l33tify(normal_string)` and `del33tify(l33t_string)` that converts a string to and from l33t sp34k, respectively. The idea is that the administration can use the `l33tify()` function to on-the-fly convert their communication to a language the kids can understand. We then pass the kids' response back through `del33tify()` so that the administrators can understand.

For example, the text from the [frontpage of the informatics department](https://www.uib.no/en/ii) looks as follows:

> With a wide range of study programs, competitive PhD positions and six recognized research groups, you have many opportunities at the Department of Informatics at UiB. Graduates from us are highly sought after in the labor market, and we prepare our graduates for highly diverse careers in ICT, finance, media and academia. Did you start your computer science studies outside of Norway? We have several options for you to be a student of our department or to start a research career with us.

Converted to l33t sp34k it looks as follows:
> With a widzor rangzor of study programs, competitivzor PhD positions and six recogniz'd research groups, joo havzor many opportunities at teh Department of Informatics at UiB. Graduates from us arzor highly sought aftxor in teh labor market, and wzor preparzor our graduates for highly diverszor careers in ICT, finance, media and academia. Did joo start joor computxor scienczor studies outsidzor of Norway? Wzor havzor several options for joo to bzor a student of our department or to start a research carexor with us.

In the regular problems we'll develop the code needed to convert short strings. In getting ahead you'll have the chance to write a program that uses this code to convert a string containing many words like the one above.

Google translate is essentially a more sophistcated version of what we're going to build in this assignment. Furthermore, computer languages, e.g., Python, do something similar internally as well. When you run your code the computer first passes through the text representation of the code (which is made for humans) and converts it to another representation that the computer can understand and execute.

As with the previous assignment you're going to work on the same code throughout the assignment. Initially the code will be short. However, as we include support for a wider range of language you'll need to add more things. Keep in mind that adding additional functions that are called from the main function is a good strategy to keep each function simple and easy to reason about. Furthermore, you'll need to take some care to figure out what's wrong when the tests don't work.

## General Guidelines
These are general guidelines that will (hopefully) help you preserve your sanity through your programming endeavours.
* Set aside time to clean up your code. The number that gets bandied around on the internet is that the best programmers spend roughly half their time cleaning up and rewriting their code, i.e., you are halfway done when you get it working.
* The sooner you start programming the longer it takes, i.e., make sure you understand the problem and have an idea for how to solve it before you get started.
* Make sure you understand the current assignment before moving on. Building on bedrock is better than building on sand.
* You should be spending a significant amount of time on [Stack Overflow](https://stackoverflow.com/) and [Google](https://duckduckgo.com/?q=how+to+python). Even experienced programmers do this.
* [PEP8](https://www.python.org/dev/peps/pep-0008/) is the official Python style guilde. Following these guidelines makes sure your code is neat and easy to read by you and others.

## Function
You'll be working in the code cell below throughout the whole assignment. As we're adding more rules you should make modifications as needed to this function. **Make sure to re-run the whole notebook by pressing the double play button at the top (restart the kernel, then re-run the whole notebook) whenever you make any changes.** This is due to two reasons:
* You need to rerun a cell with code in it whenever you make changes for the changes to take effect.
* You should re-run all tests whenever you change something to make sure your change didn't break an earlier test.
Re-running the whole notebook takes care of this automatically.

The [signature](https://en.wikipedia.org/wiki/Type_signature) of the functions you'll be working on are `l33tify(norm_speak)` and `del33tify(leet_speak)`. In both cases the parameter is a string.

When learning programming as you do now it's common to look back on the code written by past you and wonder what in the world he/she was thinking. This is good because it means present you is a stronger programmer than past you. However, it also means present you has to keep correcting the things past you wrote. Therefore, it's important that you go back and clean up your code after completing each problem.

#### Python Tip: Functions, Parameters, Arguments and Return Values
```Python
def f(a, b):
    return a + b
```
Remember that when a function is called (e.g., `f(1+2, 2*7)`) the *actual argument* values are computed and placed in the *formal parameter* variables (e.g., `a = 3`, `b = 14`). When the function reaches the `return` statement, the *return value* is computed and becomes the value of the function call.

<!-- grade_id: cell-dd49ba701655981b, points: 0, grade: false, solution: true -->
```python
dic={"elite" : "leet", "for the win" : "ftw" , "hacker" : "haxzor", "you" : "joo",
        "oh well" : "meh" , "newcomer" : "noob", "the" : "teh", "oh my god" : "zoMG!!!11!", "no way" : "nowai",
        "strong" : "stronk", "good luck" : "gl", "have fun" : "hf", "for the win": "ftw" }





def modyf(string):
    if string[-2:]=='er':
        #print(s[0:-2]+string[-2:0].replace('er','xor'))
        return string[:-2]+'xor'
    if string[-2:]=='ed':
        return string[0:-2]+'\'d'
    if string[-1:]=='e':
        return string[0:-1]+'zor'
    return string 
    

def un_modyf(string):
    if string[-3:]=='xor':
        #print(s[0:-2]+string[-2:0].replace('er','xor'))
        return string[:-3]+'er'
    if string[-2:]=='\'d':
        return string[0:-2]+'ed'
    if string[-3:]=='zor':
        return string[0:-3]+'e'
    return string 
    
def l33tify(norm_speak):
    '''convert the string norm_speak from normal speak into 
    l33t sp34k.
    
    '''
    # problem 1,1a
    leet = norm_speak.lower()
    
    
    #leet = dic[norm]
    
    for key , value in dic.items():
        leet = leet.replace(key, value)
       
    # handle capit.
    if iscapitalized(norm_speak):
        if leet=="good luck" or leet=="have fun" or leet=="for the win":
              leet = leet.upper()
        else:
            leet = capitalize(leet)
    leet = modyf(leet)
    return leet

#problem 3b
    #print(norm_speak)
    #return modyf(norm_speak)
  
    
def del33tify(leet_speak):
    '''convert the string l33t_speak from l33t sp34k into 
    normal speak.
    
    '''
    norm = leet_speak.lower() # take leet_s. na normalizate of it to lower let.
    # cre. tabele
    norm_from_leet = dict()
    for key, value in dic.items():
        norm_from_leet[value.lower()] = key  # for every key and values 
        
    for key , value in norm_from_leet.items():
        norm = norm.replace(key, value)
        
    if iscapitalized(leet_speak):
        norm = capitalize(norm)
    
    norm = un_modyf(norm)
    
    return norm
        
    
    '''  
  if iscapitalized(leet_speak):
        #if leet_speak.lower()=="gl" or leet_speak.lower()=="hf" or leet_speak.lower()=="ftw":
            #print(leet_speak)
            #return dic_o[leet_speak.lower()].upper()
        return capitalize(dic_o[leet_speak.lower()])
    return dic_o[leet_speak]
    '''


def iscapitalized(s):
    if not s:
        return False
    return s[0].isupper()
    

def capitalize(s):
    if not s:
        return ''
    
    s0=s[0].upper()
    s1=s[1:]
    return s0+s1

def endswidth(s, ending):
    '''
    Problem 3a
Implement a new function endswith(s, ending) that takes two strings as arguments and returns True if "s" ends with "ending". For example,
endswidth("hacker", "er") returns True.
endswidth("haxxor", "er") returns False.
endswidth("age", "e") returns True.
endswidth("aged", "e") returns False.
endswidth("accepted", "ed") returns True.
endswidth("accept", "ed") returns False.
    '''
    if s[-2:0]=='er':
        return s[0:-2]+s[-2:0].replace('er','xor')
    if s[-2:0]=='ed':
        return s[0:-2]+s[-2:0].replace('ed', "'d")
    if s[-1:0]=='e':
        return s[0:-1]+s[-1:0].repalce('e','zor')
    
    return s[-len(ending):]==ending
    
```

## Problem 1
The first barrier to communicating with the kids is the vocabulary. We'll need to replace some English words with the corresponding words in l33t sp34k.

### Problem 1.A
Update the `l33tify()` function such that when given the string on the left in this list it returns the one on the right (after the -> arrow). The list consists of a selection of words from the [l33t Wikipedia page](https://simple.wikipedia.org/wiki/Leet).
* `"elite"` -> `"leet"`
* `"for the win"` -> `"ftw"`
* `"hacker"` -> `"haxzor"`
* `"you"` -> `"joo"`
* `"oh well"` -> `"meh"`
* `"newcomer"` -> `"noob"`
* `"the"` -> `"teh"`
* `"oh my god"` -> `"zoMG!!!11!"`
* `"no way"` -> `"nowai"`
* `"strong"` -> `"stronk"`
* `"good luck"` -> `"gl"`
* `"have fun"` -> `"hf"`

```python
# Python hint: Replacing words
s1 = 'foobar'
s2 = s1.replace('foo', '123') # note the variable assignment
print('s1', s1) # prints 'foobar'. s1 remains unchanged.
print('s2', s2) # prints '123bar'. s2 contains the replacement.

# you can chain multiple replacements
s = 'foobar'
s = s.replace('foo', '123') # assignment back to the same variable
s = s.replace('bar', '456')
print('s', s)
```

Output:
```
s1 foobar
s2 123bar
s 123456
```

<!-- grade_id: cell-e1f7c8f2bae79f2e, points: 2, grade: true, solution: false -->
```python
def check(norm_speak, leet_speak, lfn=True, nfl=True):
    if lfn:
        leet_from_norm = l33tify(norm_speak)
        assert leet_from_norm == leet_speak, f'error when converting norm-speak to l33t-sp34k: expected {leet_speak}, got {leet_from_norm}'
    
    if nfl:
        norm_from_leet = del33tify(leet_speak)
        assert norm_from_leet == norm_speak, f'error when converting l33t-sp34k to norm-speak: expected {norm_speak}, got {norm_from_leet}'
    
def testProblem1a():
    print("Problem 1a:", end='')
    check("elite", "leet", nfl=False)
    check("for the win", "ftw", nfl=False)
    check("hacker", "haxzor", nfl=False)
    check("you", "joo", nfl=False)
    check("oh well", "meh", nfl=False)
    check("newcomer", "noob", nfl=False)
    check("the", "teh", nfl=False)
    check("oh my god", "zoMG!!!11!", nfl=False)
    check("no way", "nowai", nfl=False)
    check("strong", "stronk", nfl=False)
    check("good luck", "gl", nfl=False)
    check("have fun", "hf", nfl=False)
    print("    tests pass")

testProblem1a()
```

Output:
```
Problem 1a:    tests pass
```

### Problem 1.B

Now update the `del33tify()` function accordingly, i.e., it should make the opposite conversion from `l33tify()` such that `del33tify(l33tify(s)) == s` is `True`.

<!-- grade_id: cell-45eb2240cd28593e, points: 2, grade: true, solution: false -->
```python
def testProblem1b():
    print("Problem 1b:", end='')
    check("elite", "leet")
    check("for the win", "ftw")
    check("hacker", "haxzor")
    check("you", "joo")
    check("oh well", "meh")
    check("newcomer", "noob")
    check("the", "teh")
    check("oh my god", "zoMG!!!11!")
    check("no way", "nowai")
    check("strong", "stronk")
    check("good luck", "gl")
    check("have fun", "hf")
    print("    tests pass")

testProblem1b()
testProblem1a()
```

Output:
```
Problem 1b:    tests pass
Problem 1a:    tests pass
```

## Problem 2

We were playing fast and loose with the capitalization in the previous problem and the administrators are not happy. They insist on consistent rules for how words should be capitalized when converting to and from l33t sp34k.

### Problem 2a

Write a function `iscapitalized(s)` that takes a string as its single argument and returns `True` if the first letter of `"s"` is capitalized, e.g., `"l33t"` -> `False` and `"L33t"` -> `True`.

#### Python hint: Defining new functions
You define new functions next to each other like this:
```Python
def f():
    print('running function f')
    
def g():
    print('running function g')

f() # running function f
g() # running function g
```

If you define `g()` inside of `f()`, `g()` is not visible outside `f()`, i.e., the below won't work. Note the indentation.
```Python
def f():
    print('running function f')
    
    def g():
        print('running function g')

f() # running function f
g() # causes an error since g() is only visible inside f()
```

```python
# Python hint: Substrings
s = 'Foobar'
s0 = s[0] # the [i] syntax allows you to extract the i-th character
print(s0) 

s1 = s[1] # it works for any value of 0 <= i < len(s)
print(s1)

s0_2 = s[0:2] # the [i1:i2] syntax allows you to extract a substring
print(s0_2)

s3_5 = s[3:5] # works for any values 0 <= i1 <= i2 < len(s)
print(s3_5)
```

Output:
```
F
o
Fo
ba
```

```python
# Python hint: isupper
s1 = 'F'
s2 = 'f'
print(s1.isupper()) # True
print(s2.isupper()) # False
```

Output:
```
True
False
```

<!-- grade_id: cell-adf572ed466ef5d7, points: 0, grade: true, solution: false -->
```python
def assertEqual(a, b, tol=0):
    assert a == b, f'expected {b}, got {a}'
    
def testProblem2a():
    print("Problem 2a:", end='')
    assertEqual(iscapitalized("l33t"), False)
    assertEqual(iscapitalized("L33t"), True)
    assertEqual(iscapitalized("l33T"), False)
    assertEqual(iscapitalized("L33T"), True)
    assertEqual(iscapitalized("l33t Haxzors"), False)
    assertEqual(iscapitalized("L33t Haxzors"), True)
    assertEqual(iscapitalized("1F33t"), False)
    assertEqual(iscapitalized(""), False)
    print("    tests pass")

testProblem2a()
```

Output:
```
Problem 2a:    tests pass
```

### Problem 2b
Write a function `capitalize(s)` that takes a string as its single argument and returns a new string that is exctly the same except that the first character is capitalized, e.g., `"l33t"` -> `"L33t"`.

<!-- grade_id: cell-d2e50a737f9615d8, points: 0, grade: true, solution: false -->
```python
def testProblem2b():
    print("Problem 2b:", end='')
    assertEqual(capitalize("l33t"), "L33t")
    assertEqual(capitalize("L33t"), "L33t")
    assertEqual(capitalize("l33t Haxzors"), "L33t Haxzors")
    assertEqual(capitalize("L33t Haxzors"), "L33t Haxzors")
    assertEqual(capitalize(""), "")
    print("    tests pass")

testProblem2b()
testProblem2a()
```

Output:
```
Problem 2b:    tests pass
Problem 2a:    tests pass
```

```python
# Python hint: Concatenating strings
s1 = 'foo'
s2 = 'bar'

s3 = s1 + s2
print('s3', s) # foobar

s4 = s1[1:2] + s2[0:2] # o + ba
print('s4', s4) # oba
```

Output:
```
s3 Foobar
s4 oba
```

```python
# Python hint: String length
s = 'foobar'
print('len of s', len(s)) # s is 6 characters long
```

Output:
```
len of s 6
```

### Problam 2c
Update the `l33tify()` function with the following rule for capitalization. Use the functions `iscapitalized()` and `capitalize()` that you wrote.
* If the first letter in normal speak is capitalized then the first letter in l33t sp34k should also be capitalized. For example, `"Elite"` converts to `"L33t"` and `"newcomer"` converts to `"noob"`.
* Note that the `"MG`" in `"ZoMG!!!11!"` should always be capitalized. The `"Z"` is capitalized only if the first letter in norm speak is capitalized, e.g., `"oh my god"` converts to `"zoMG!!!11!"` and `"Oh my god"` converts to `"ZoMG!!!11!"`.

<!-- grade_id: cell-b72ca823db5436a5, points: 2, grade: true, solution: false -->
```python
def testProblem2c(verbose=True, nfl=False):
    if verbose:
        print("Problem 2c:", end='')
    check("elite", "leet", nfl=nfl)
    check("hacker", "haxzor", nfl=nfl)
    check("you", "joo", nfl=nfl)
    check("oh well", "meh", nfl=nfl)
    check("newcomer", "noob", nfl=nfl)
    check("the", "teh", nfl=nfl)
    check("no way", "nowai", nfl=nfl)
    check("strong", "stronk", nfl=nfl)
    check("oh my god", "zoMG!!!11!", nfl=nfl)
    check("Elite", "Leet", nfl=nfl)
    check("Hacker", "Haxzor", nfl=nfl)
    check("You", "Joo", nfl=nfl)
    check("Oh well", "Meh", nfl=nfl)
    check("Newcomer", "Noob", nfl=nfl)
    check("The", "Teh", nfl=nfl)
    check("No way", "Nowai", nfl=nfl)
    check("Strong", "Stronk", nfl=nfl) 
    check("Oh my god", "ZoMG!!!11!", nfl=nfl)
    if verbose:
        print("    tests pass")
    
testProblem2c()
testProblem2b()
testProblem2a()
```

Output:
```
Problem 2c:    tests pass
Problem 2b:    tests pass
Problem 2a:    tests pass
```

### Problem 2d
 The capitalization rules are noy yet completely consistent. Add the following two rules to the `l33tify()` function.
* For `"ftw"`, `"GL"` and `"HF"` all letters should be capitalized if the first letter of `For the win"`, `"Good luck"` or `"Have fun"` is capitalized, e.g.,
    * `"For the win"` -> `"FTW"` and `"for the win"` -> `"ftw"`
    * `"Good luck"` -> `"GL"` and `"good luck"` -> `"gl"`
    * `"Have fun"` -> `"HF"` and `"have fun"` -> `"hf"`

#### Hey, listen! if and elif
Use `if` with `elif` and/or `else` to select one of several alternatives:
```python
s = "foo"
v = 10
if s == "foo" and v < 8: # run the indented code if v is smaller than 8
    print("v is less than 8")
     # do more stuff if you want, e.g. 'return' or 'a = b', or something
elif s == "foo" and v < 12: # run if v is not smaller than 8 but is smaller than 12
    print("v is less than 12 and not smaller than 8")
    # do more stuff if you want, e.g. 'return' or 'a = b', or something
```

In general, it is often better to use `elif`/`else` even if all your `if`s end with `return`, since that makes it clearer to the reader that we are selecting one of several alternatives:
```python
    # ....
    if a == b:
        return "equal"
    else:  
        return "different"
```

<!-- grade_id: cell-2e0d2298149a955a, points: 2, grade: true, solution: false -->
```python
def testProblem2d(verbose=True, nfl=False):
    if verbose:
        print("Problem 2d:", end='')
    check("for the win", "ftw", nfl=nfl)
    check("oh my god", "zoMG!!!11!", nfl=nfl)
    check("good luck", "gl", nfl=nfl)
    check("have fun", "hf", nfl=nfl)
    check("For the win", "FTW", nfl=nfl)
    check("Good luck", "GL", nfl=nfl)
    check("Have fun", "HF", nfl=nfl)
    if verbose:
        print("    tests pass")

testProblem2d()
testProblem2c()
testProblem2b()
testProblem2a()
```

Output:
```
Problem 2d:<ipython-input-224-cdcf68bab0a4> in testProblem2d(verbose, nfl)
      6     check("good luck", "gl", nfl=nfl)
      7     check("have fun", "hf", nfl=nfl)
----> 8     check("For the win", "FTW", nfl=nfl)
      9     check("Good luck", "GL", nfl=nfl)
     10     check("Have fun", "HF", nfl=nfl)

<ipython-input-209-4c150ba0a3af> in check(norm_speak, leet_speak, lfn, nfl)
      2     if lfn:
      3         leet_from_norm = l33tify(norm_speak)
----> 4         assert leet_from_norm == leet_speak, f'error when converting norm-speak to l33t-sp34k: expected {leet_speak}, got {leet_from_norm}'
      5 
      6     if nfl:

AssertionError: error when converting norm-speak to l33t-sp34k: expected FTW, got Ftw
```

### Problem 2e

The administrators are happy with your progress. However, they still can't understand what the kids are saying. Update the `del33tify()` function to take into account the things we added to the `l33tify()` function. Remember that it should make the opposite conversion from `l33tify()` such that `del33tify(l33tify(s)) == s` is `True`.

<!-- grade_id: cell-dcacfb257b0ca67f, points: 0, grade: true, solution: false -->
```python
def testProblem2e():
    print("Problem 2e:", end='')
    testProblem2d(nfl=True, verbose=False)
    testProblem2c(nfl=True, verbose=False)
    print("    tests pass")

testProblem2e()
testProblem2d()
testProblem2c()
testProblem2b()
testProblem2a()
```

Output:
```
Problem 2e:<ipython-input-215-cdcf68bab0a4> in testProblem2d(verbose, nfl)
      6     check("good luck", "gl", nfl=nfl)
      7     check("have fun", "hf", nfl=nfl)
----> 8     check("For the win", "FTW", nfl=nfl)
      9     check("Good luck", "GL", nfl=nfl)
     10     check("Have fun", "HF", nfl=nfl)

<ipython-input-209-4c150ba0a3af> in check(norm_speak, leet_speak, lfn, nfl)
      2     if lfn:
      3         leet_from_norm = l33tify(norm_speak)
----> 4         assert leet_from_norm == leet_speak, f'error when converting norm-speak to l33t-sp34k: expected {leet_speak}, got {leet_from_norm}'
      5 
      6     if nfl:

AssertionError: error when converting norm-speak to l33t-sp34k: expected FTW, got Ftw
```

## Problem 3
Nice work! We'll soon be able to communicate with the kids. The next thing we're going to handle is the most common l33t sp34k word endings. Specifically, words in l33t sp34k often end with -xor, -zor, -age, -'d.

### Problem 3a
Implement a new function `endswith(s, ending)` that takes two strings as arguments and returns `True` if `"s"` ends with `"ending"`. For example,
* `endswidth("hacker", "er")` returns `True`.
* `endswidth("haxxor", "er")` returns `False`.
* `endswidth("age", "e")` returns `True`.
* `endswidth("aged", "e")` returns `False`.
* `endswidth("accepted", "ed")` returns `True`.
* `endswidth("accept", "ed")` returns `False`.

Do you remember how to extract a substring? :)

<!-- grade_id: cell-014f0c1c121f6d93, points: 0, grade: true, solution: false -->
```python
def testProblem3a():
    print("Problem 3a:", end='')
    assertEqual(endswidth("hacker", "er"), True)
    assertEqual(endswidth("haxxor", "er"), False)
    assertEqual(endswidth("age", "e"), True)
    assertEqual(endswidth("aged", "e"), False)
    assertEqual(endswidth("accepted", "ed"), True)
    assertEqual(endswidth("accept", "ed"), False)
    assertEqual(endswidth("accept", "accept"), True)
    assertEqual(endswidth("accept", "faccept"), False)
    assertEqual(endswidth("", ""), True)
    assertEqual(endswidth(" ", " "), True)
    print("    tests pass")
    
testProblem3a()
```

Output:
```
Problem 3a:    tests pass
```

### Problem 3b
Implement the following rules:
* Words ending with `"er"` in normal speak should instead end with `"xor"` in l33t sp34k.
* Words ending with `"e"` in normal speak should instead end with `"zor"` in l33t sp34k.
* Words ending with `"ed"` in normal speak should instead end with `"'d"` in l33t sp34k.

These rules should be applied after those added in problem 1. Note that `s.replace(s1, s2)` replaces all ocurrences (not only the ending) by default. You need to both extract a substring and concatenate strings here.

<!-- grade_id: cell-f61b3acfacd73561, points: 3, grade: true, solution: false -->
```python
def testProblem3b(verbose=True, nfl=False):
    if verbose:
        print("Problem 3b:", end='')
    check("weaker", "weakxor", nfl=nfl)
    check("weak", "weak", nfl=nfl)
    check("age", "agzor", nfl=nfl)
    check("based", "bas'd", nfl=nfl)
    check("knowledge", "knowledgzor", nfl=nfl)
    check("understood", "understood", nfl=nfl)
    check("referred", "referr'd", nfl=nfl)
    if verbose:
        print("    tests pass")

testProblem3b()
testProblem3a()
```

Output:
```
Problem 3b:    tests pass
Problem 3a:    tests pass
```

### Problem 3c
All right. Final problem of the assignment! Update the `del33tify()` function to take the word endings into account. Remember that it should make the opposite conversion from `l33tify()` such that `del33tify(l33tify(s)) == s` is `True`.

<!-- grade_id: cell-39bcf2706e6ab203, points: 3, grade: true, solution: false -->
```python
def testProblem3c():
    print("Problem 3c:", end='')
    testProblem3b(verbose=False, nfl=True)
    print("    tests pass")

testProblem3c()
testProblem3b()
testProblem3a()
```

Output:
```
Problem 3c:    tests pass
Problem 3b:    tests pass
Problem 3a:    tests pass
```

## Getting Ahead
This week we've focused on dealing with strings. We've also worked with separating functionality via functions. Congratulations on completing it! Processing strings comes up in a wide range of practical applications, e.g., extracting information from strings given by users. It's easy to imagine that spell correction systems does something similar to this.

So far our system has only dealt with a single word at a time. To translate entire texts we need to be able to translate strings containing many words to and from l33t sp34k. In this section we're going to implement this. Specifically, write two new functions `l33tify_many(norm_speak)` and `del33tify_many(leet_speak)` that both take a string containing many words as its only argument and converts the whole thing to and from l33t sp34k.
* We don't care about capitalization here. For example, it's fine to convert everything to lower case.
* The word replacement comes first when converting to l33t sp34k and last when converting back. Word endings are handled after replacing words when converting to l33t sp34k and before replacing words when converting back.

This part is meant as preparation for the coming course material and includes materiel not yet covered in the course. Then again, you never have to catch up if you're always getting ahead. This part requires doing more learning on your own. You're in very good shape if you can complete it :)

```python
# Python hint: Splitting a string and iterating over the substrings
s = 'foo bar'
l = s.split(' ') # returns a list ['foo', 'bar']
for w in l: # iterate over the words in l
    print(w)

for i in range(len(l)): # same as above but you get the index
    l[i] = l[i].upper() # assign a new string to element i of the list
    print(l[i], 'is at index', i)
    
# Python hint: Joining strings
s2 = ''.join(l) # appends the strings in l to the empty string ('')
print('s2', s2) # note that s2 is a single string
```

Output:
```
foo
bar
FOO is at index 0
BAR is at index 1
s2 FOOBAR
```

<!-- grade_id: cell-89054e03a65466b2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


def l33tify_many(norm_speak):
    '''convert the string norm_speak containing several words 
    from normal speak into l33t sp34k.
    
    '''
    # YOUR CODE HERE
    
    leet_speak = norm_speak.lower().split(' ')
    for i in range(len(leet_speak)):
        leet_speak[i] = l33tify(leet_speak[i])
    return leet_speak
        
        
def del33tify_many(leet_speak):
    '''convert the string l33t_speak containing several words 
    from l33t sp34k into normal speak.
    
    '''
    # YOUR CODE HERE
    
    for i in range(len(leet_speak)):
        leet_speak[i] = ' ' +del33tify(leet_speak[i])
    return ''.join(leet_speak).strip()
        
```

<!-- grade_id: cell-ef1edcdd99fb913f, points: 0, grade: true, solution: false -->
```python
def check_many(s):
    leet_speak = l33tify_many(s)
    norm_speak = del33tify_many(leet_speak)
    print('Original text:')
    print('---------------')
    print(s)
    print()
    print('Converted to l33t sp34k:')
    print('---------------')
    print(leet_speak)
    print()
    print('Converted back to norm speak:')
    print('---------------')
    print(norm_speak)
    print()
    
    # make lower case, split by word and strip whitespace
    s_words = [w.strip() for w in s.lower().split(' ')]
    norm_words = [w.strip() for w in norm_speak.lower().split(' ')]
    
    # check that the number of words is correct
    assert len(s_words) == len(norm_words), \
        f'incorrect length after converting back to norm speak. expected {len(s_words)} words, got {len(norm_words)} words'

    # compare word by word
    for i, (s_word, norm_word) in enumerate(zip(s_words, norm_words)):
        assert s_word == norm_word, f'incorrect translation for word {i}. expected {s_word}, got {norm_word}'

# Norwegian
# https://www.uib.no/ii
s = '''
Informatikk tilsvarer computer science p engelsk, og er vitenskapen som ligger bak enhver programvare, fra enkle dataprogram, mobilapper og spill, til store IT-systemer, databaser og nettportaler.

Informatikere er svrt ettertraktet i arbeidsmarkedet fordi samfunnet er blitt avhengig av at informasjonsteknologien fungerer. Med et stort utvalg studieprogrammer, konkurransedyktige ph.d.-stillinger og seks anerkjente forskergrupper, har du mange muligheter p institutt for informatikk ved UiB.
'''
check_many(s)

# English
# https://www.uib.no/en/ii
s = '''
With a wide range of study programs, competitive PhD positions and six recognized research groups, you have many opportunities at the Department of Informatics at UiB. Graduates from us are highly sought after in the labor market, and we prepare our graduates for highly diverse careers in ICT, finance, media and academia. Did you start your computer science studies outside of Norway? We have several options for you to be a student of our department or to start a research career with us.

Learn more about our bachelor, master and PhD programs or our exchange programs.

Here you can find admission information on how to apply to the University of Bergen.
'''
check_many(s)

print('   tests pass')
```

Output:
```
Original text:
---------------

Informatikk tilsvarer computer science p engelsk, og er vitenskapen som ligger bak enhver programvare, fra enkle dataprogram, mobilapper og spill, til store IT-systemer, databaser og nettportaler.

Informatikere er svrt ettertraktet i arbeidsmarkedet fordi samfunnet er blitt avhengig av at informasjonsteknologien fungerer. Med et stort utvalg studieprogrammer, konkurransedyktige ph.d.-stillinger og seks anerkjente forskergrupper, har du mange muligheter p institutt for informatikk ved UiB.


Converted to l33t sp34k:
---------------
[' \ninformatikk', ' tilsvarer', ' computer', ' science', ' p', ' engelsk,', ' og', ' er', ' vitenskapen', ' som', ' ligger', ' bak', ' enhver', ' programvare,', ' fra', ' enkle', ' dataprogram,', ' mobilapper', ' og', ' spill,', ' til', ' store', ' it-systemer,', ' databaser', ' og', ' nettportaler.\n\ninformatikere', ' er', ' svrt', ' ettertraktet', ' i', ' arbeidsmarkedet', ' fordi', ' samfunnet', ' er', ' blitt', ' avhengi
```

```python

```